# Statystyka matematyczna - wiczenia laboratoryjne 2023/2024

Ten notatnik zalicza si do grupy zestaw贸w zada, na podstawie kt贸rych odbywa si zaliczenie wicze i podlega zwrotowi do oceny w ustalonym na zajciach terminie.

Uwagi i wytyczne og贸lne dotyczce uzupeniania notatnika:
- Podczas wykonywania zada nale偶y korzysta wycznie z pakiet贸w zaimportowanych na pocztku notatnika oraz z pakiet贸w wchodzcych w skad standardowej biblioteki Pythona, kt贸re mo偶na zaimportowa samodzielnie we wskazanej kom贸rce notatnika.
- Swoje rozwizania nale偶y wprowadza wycznie w miejce nastpujcych fragment贸w kodu:<br/> `# YOUR CODE HERE`<br/> `raise NotImplementedError()`<br/> 
a odpowiedzi tekstowe w kom贸rkach oznaczonych hasem:<br/> 
`YOUR ANSWER HERE`<br/> 
Nie nale偶y w 偶aden spos贸b modyfikowa pozostaych fragment贸w kodu oraz innych element贸w notatnika, w szczeg贸lnoci dodawa lub usuwa kom贸rek oraz zmienia nazwy pliku.
- Otrzymywane wyniki i odpowiedzi maj by rezultatem wykonania napisanego kodu.
- Zadanie nale偶y wykona w taki spos贸b, aby podczas wykonywania kodu nie zostay wywietlone 偶adne ostrze偶enia.
- Zawarte w notatniku automatyczne testy maj charakter pogldowy. Dotycz one wybranych aspekt贸w zada i maj za zadanie wyapa podstawowe bdy. Przejcie wszystkich test贸w nie oznacza, 偶e zadanie jest wykonane w caoci poprawnie.

Uwagi i wytyczne og贸lne dotyczce wizualizacji wynik贸w:
- Wszystkie wykresy powinny by wykonane w jednolitym, przejrzystym i czytelnym stylu, posiada odpowiednio dobrane proporcje i zakresy wartoci osi.
- Wykresy oraz ich osie powinny mie nadane tytuy. Je偶eli w obrbie figury znajduje si wicej ni偶 jeden wykres to figura r贸wnie偶 powinna mie nadany tytu. 
- Figury powinny mie ustawione biae to, tak, aby niezale偶nie od ustawie notatnika wszystkie elementy wykres贸w byy dobrze widoczne (domylnie to jest prze藕roczyste co mo偶e powodowa problemy w notatnikach z ustawionym ciemnym tem).
- Rozmiar poziomy figur nie powinien przekracza 20 cali.

Przed odesaniem zestawu zada do oceny prosz uzupeni kom贸rk z danymi autor贸w rozwizania (nazwa zespou oraz imiona, nazwiska i numery indeks贸w czonk贸w zespou) oraz upewni si, 偶e notatnik dziaa zgodnie z oczekiwaniami. W tym celu nale偶y skorzysta z opcji **Restart Kernel and Run All Cells...** dostpnej na g贸rnej belce notatnika pod symbolem $\blacktriangleright\blacktriangleright$.

Nazwa zespou:
Czonkowie:

---

# Zestaw zada 6: Estymacja parametryczna

```python
import numpy as np
import pandas as pd
import scipy as sp
import matplotlib as mpl
import matplotlib.pyplot as plt
import seaborn as sns
```

```python
# Miejsce do importu pakiet贸w wchodzcych w skad standardowej biblioteki Pythona oraz ustawienie opcji wykorzystywanych pakiet贸w
from scipy.stats import t
from scipy.stats import norm
import math
```

### Dane do zada

W celu wygenerowania danych wykorzystywanych w zawartych w notatniku kom贸rkach testowych wykonaj obie poni偶sze kom贸rki.

```python
# Dane do test贸w 1
test_data_1 = pd.DataFrame(data=sp.stats.norm.rvs(loc=5, scale=0.2, size=20, random_state=7), columns=["X"])
test_data_1.head()
```

Output:
```
X
0  5.338105
1  4.906813
2  5.006564
3  5.081503
4  4.842215
```

```python
# Dane do test贸w 2
x = sp.stats.uniform.rvs(loc=-2, scale=10, size=250, random_state=34)
y = 2*x - 5 + sp.stats.norm.rvs(loc=0, scale=2, size=250, random_state=13)
test_data_2 = pd.DataFrame(data=np.array([x, y]).T, columns=["X", "Y"])
test_data_2.head()
```

Output:
```
X         Y
0 -1.614383 -9.653548
1  5.801005  8.109542
2 -1.072962 -7.234931
3  4.328927  4.561478
4 -1.861092 -6.031980
```

### Zadanie 1: Estymacja wartoci oczekiwanej [8 pkt]

Przygotuj funkcj `mean_estimation()`, kt贸ra bdzie dokonywaa estymacji przedziaowej wartoci oczekiwanej dla danych wejciowych w postaci szeregu szczeg贸owego zgodnie ze schematem z zaczonego do notatnika zestawu wzor贸w.

Opr贸cz zwracanych wartoci granic przedziau funkcja powinna wywietla nastpujcy komunikat:

` nale偶y do przedziau [X, Y] przy zao偶eniu poziomu ufnoci 1- = Z`

gdzie X, Y i Z s automatycznie uzupeniane przez funkcj, a X i Y dodatkowo sformatowane w taki spos贸b, 偶eby wywietlay si z dokadnoci do 4 miejsc po przecinku.

<!-- grade_id: cell-3190cac551203c3a, points: 0, grade: false, solution: true -->
```python
from scipy.stats import norm
from scipy.stats import t
def mean_estimation(data, alpha, population_std="unknown"):
    
    
    """
    Parameters
    -------
    data: DataFrame
        Tablica zawierajca domylny indeks i pojedyncz kolumn "X" z wartociami pr贸by losowej,
    alpha: float
        Warto okrelajca poziom istotnoci.
    population_std: str or float
        Je偶eli odchylenie standardowe populacji nie jest znane to parametr przyjmuje warto "unknown",
        w przeciwnym wypadku jest to warto okrelajca odchylenie standardowe populacji.
    
    Returns
    -------
    mean_low: float
        Dolna granica wyliczonego przedziau ufnoci.
    mean_high: float
        G贸rna granica wyliczonego przedziau ufnoci.
    """
    n = len(data)
    X_mean = 1/n*np.sum(data['X'])
    if population_std == "unknown":
        t_s = t.ppf(1 - (alpha / 2), df=n-1) #rozkad t-studenta 
        s = 1 / (n-1) * np.sum((data['X'] - X_mean) ** 2)
        margin = t_s * np.sqrt(s) / np.sqrt(n)
    else:
        t_s = norm.ppf(1 - alpha / 2)
        margin = t_s * population_std / np.sqrt(n)
    
    mean_low = X_mean - margin
    mean_high = X_mean + margin
    
    print(f" nale偶y do przedziau [{mean_low:.4f}, {mean_high:.4f}] przy zao偶eniu poziomu ufnoci 1-{alpha} = {1-alpha}")

    return mean_low, mean_high
    # YOUR CODE HERE
    #raise NotImplementedError()
```

<!-- grade_id: cell-f856aaac1658539a, points: 8, grade: true, solution: false -->
```python
### Kom贸rka testowa

# Test na test_data_1 przy zao偶eniu, 偶e znane jest odchylenie standardowe populacji, z kt贸rej zostaa pobrana pr贸bka losowa
assert(np.all(np.isclose(mean_estimation(test_data_1, 0.1, population_std=0.2), (4.905619719495418, 5.05273990041144))))
print()
assert(np.all(np.isclose(mean_estimation(test_data_1, 0.05, population_std=0.2), (4.891527555895771, 5.066832064011087))))
print()
assert(np.all(np.isclose(mean_estimation(test_data_1, 0.02, population_std=0.2), (4.875142370240081, 5.083217249666777))))
print()

# # Test na test_data_1 przy zao偶eniu, 偶e nie jest znane odchylenie standardowe populacji, z kt贸rej zostaa pobrana pr贸bka losowa
assert(np.all(np.isclose(mean_estimation(test_data_1, 0.1), (4.913356369231727, 5.045003250675131))))
print()
assert(np.all(np.isclose(mean_estimation(test_data_1, 0.05), (4.899504007229391, 5.058855612677467))))
print()
assert(np.all(np.isclose(mean_estimation(test_data_1, 0.02), (4.882508507351398, 5.07585111255546))))
```

Output:
```
nale偶y do przedziau [4.9056, 5.0527] przy zao偶eniu poziomu ufnoci 1-0.1 = 0.9

 nale偶y do przedziau [4.8915, 5.0668] przy zao偶eniu poziomu ufnoci 1-0.05 = 0.95

 nale偶y do przedziau [4.8751, 5.0832] przy zao偶eniu poziomu ufnoci 1-0.02 = 0.98

 nale偶y do przedziau [4.9134, 5.0450] przy zao偶eniu poziomu ufnoci 1-0.1 = 0.9

 nale偶y do przedziau [4.8995, 5.0589] przy zao偶eniu poziomu ufnoci 1-0.05 = 0.95

 nale偶y do przedziau [4.8825, 5.0759] przy zao偶eniu poziomu ufnoci 1-0.02 = 0.98
```

### Zadanie 2: Estymacja wariancji [8 pkt]

Przygotuj funkcj `variance_estimation()`, kt贸ra bdzie dokonywaa estymacji przedziaowej wariancji dla danych wejciowych w postaci szeregu szczeg贸owego zgodnie ze schematem z zaczonego do notatnika zestawu wzor贸w.

Opr贸cz zwracanych wartoci granic przedziau funkcja powinna wywietla nastpujcy komunikat:

`^2  nale偶y do przedziau [X, Y] przy zao偶eniu poziomu ufnoci 1- = Z`

gdzie X, Y i Z s automatycznie uzupeniane przez funkcj, a X i Y dodatkowo sformatowane w taki spos贸b, 偶eby wywietlay si z dokadnoci do 4 miejsc po przecinku.

<!-- grade_id: cell-79a02915c858343b, points: 0, grade: false, solution: true -->
```python
def variance_estimation(data, alpha, population_mean="unknown"):
    """
    Parameters
    -------
    data: DataFrame
        Tablica zawierajca domylny indeks i pojedyncz kolumn "X" z wartociami pr贸by losowej,
    alpha: float
        Warto okrelajca poziom istotnoci.
    population_mean: str or float
        Je偶eli warto oczekiwana populacji nie jest znane to parametr przyjmuje warto "unknown",
        w przeciwnym wypadku jest to warto okrelajca warto oczekiwan populacji.

    Returns
    -------
    var_low: float
        Dolna granica wyliczonego przedziau ufnoci.
    var_high: float
        G贸rna granica wyliczonego przedziau ufnoci.
    """
    n = data.shape[0]
    X = data[data.columns[0]]

    if population_mean=="unknown":
        S2 = X.var(ddof=0)
    else:
        S2 = 1/n * sum((xi - population_mean)**2 for xi in X)

    var_low = n * S2 / sp.stats.chi2.ppf(1-alpha/2, n-1)
    var_high = n * S2 / sp.stats.chi2.ppf(alpha/2, n-1)

    print("^2 nale偶y do przedziau [{:.4f}, {:.4f}] przy zao偶eniu poziomu ufnoci 1- = {}".format(var_low, var_high, 1-alpha))

    return var_low, var_high
```

<!-- grade_id: cell-590df950d841e5e9, points: 8, grade: true, solution: false -->
```python
### Kom贸rka testowa

# Test na test_data_1 przy zao偶eniu, 偶e znane jest odchylenie standardowe populacji, z kt贸rej zostaa pobrana pr贸bka losowa
assert(np.all(np.isclose(variance_estimation(test_data_1, 0.1, population_mean=5), (0.01855573624556403, 0.05528660848892888))));
print()
assert(np.all(np.isclose(variance_estimation(test_data_1, 0.05, population_mean=5), (0.017025745016269588, 0.06280068548350207))));
print()
assert(np.all(np.isclose(variance_estimation(test_data_1, 0.02, population_mean=5), (0.015455150810030724, 0.07328116757245808))));
print()

# Test na test_data_1 przy zao偶eniu, 偶e nie jest znane odchylenie standardowe populacji, z kt贸rej zostaa pobrana pr贸bka losowa
assert(np.all(np.isclose(variance_estimation(test_data_1, 0.1), (0.018268125369359626, 0.05442967510189219))));
print()
assert(np.all(np.isclose(variance_estimation(test_data_1, 0.05), (0.0167618487538223, 0.06182728513230527))));
print()
assert(np.all(np.isclose(variance_estimation(test_data_1, 0.02), (0.015215598500840796, 0.07214532146342346))));
```

Output:
```
^2 nale偶y do przedziau [0.0186, 0.0553] przy zao偶eniu poziomu ufnoci 1- = 0.9

^2 nale偶y do przedziau [0.0170, 0.0628] przy zao偶eniu poziomu ufnoci 1- = 0.95

^2 nale偶y do przedziau [0.0155, 0.0733] przy zao偶eniu poziomu ufnoci 1- = 0.98

^2 nale偶y do przedziau [0.0183, 0.0544] przy zao偶eniu poziomu ufnoci 1- = 0.9

^2 nale偶y do przedziau [0.0168, 0.0618] przy zao偶eniu poziomu ufnoci 1- = 0.95

^2 nale偶y do przedziau [0.0152, 0.0721] przy zao偶eniu poziomu ufnoci 1- = 0.98
```

### Zadanie 3: Estymacja wartoci wsp贸czynnika korelacji liniowej [6 pkt]

Przygotuj funkcj `correlation_estimation()`, kt贸ra bdzie dokonywaa estymacji przedziaowej wartoci wsp贸czynnika korelacji liniowej Pearsona dla danych wejciowych w postaci szeregu szczeg贸owego zgodnie ze schematem z zaczonego do notatnika zestawu wzor贸w.

Opr贸cz zwracanych wartoci granic przedziau funkcja powinna wywietla nastpujcy komunikat:

`r nale偶y do przedziau [X, Y] przy zao偶eniu poziomu ufnoci 1- = Z`

gdzie X, Y i Z s automatycznie uzupeniane przez funkcj, a X i Y dodatkowo sformatowane w taki spos贸b, 偶eby wywietlay si z dokadnoci do 4 miejsc po przecinku.

<!-- grade_id: cell-e3388b3e7966ada6, points: 0, grade: false, solution: true -->
```python
def correlation_estimation(data, alpha):
    """
    Parameters
    -------
    data: DataFrame
        Tablica zawierajca domylny indeks i dwie kolumny "X" i "Y" z wynikami pr贸by losowej.
    alpha: float
        Warto okrelajca poziom istotnoci.
    Returns
    -------
    r_corr_low: float
        Dolna granica wyliczonego przedziau ufnoci.
    r_corr_high: float
        G贸rna granica wyliczonego przedziau ufnoci.
    """
    # YOUR CODE HERE
    mean_x = 0
    mean_y = 0
    n = len(data)
    for i in data.iloc[:,0]:
        mean_x+=i
    mean_x/=n
    for i in data.iloc[:,1]:
        mean_y+=i
    mean_y/=n
    r = 0
    for i in range(n):
        r+=(data.iloc[i,0]-mean_x)*(data.iloc[i,1]-mean_y)
    temp = 0
    temp1= 0
    temp2= 0
    for i in range(n):
        temp1+=(data.iloc[i,0]-mean_x)**2
        temp2+=(data.iloc[i,1]-mean_y)**2
    temp1 = math.sqrt(temp1)
    temp2 = math.sqrt(temp2)
    temp = temp1 * temp2
    r/=temp
    Z = 1 - alpha
    u = norm.ppf(1-alpha/2)
    r_corr_low = r-u*((1-r**2)/math.sqrt(n))
    r_corr_high = r+u*((1-r**2)/math.sqrt(n))
    print(f'r nalezy do przedzialu [{r_corr_low:.4},{r_corr_high:.4}] przy zao偶eniu poziomu ufnoci 1 - 伪 = {Z:.4}')
    return (r_corr_low, r_corr_high)
```

<!-- grade_id: cell-2f4af63f6b0b79b7, points: 6, grade: true, solution: false -->
```python
### Kom贸rka testowa
assert np.all(np.isclose(correlation_estimation(test_data_2, 0.05), (0.9494696249316307, 0.9692185148958979)))
assert np.all(np.isclose(correlation_estimation(test_data_2, 0.02), (0.9476237555239514, 0.9710643843035772)))
```

Output:
```
r nalezy do przedzialu [0.9495,0.9692] przy zao偶eniu poziomu ufnoci 1 - 伪 = 0.95
r nalezy do przedzialu [0.9476,0.9711] przy zao偶eniu poziomu ufnoci 1 - 伪 = 0.98
```

### Zadanie 4: Estymacja wsp贸czynnik贸w r贸wnania regresji liniowej [8 pkt]

Przygotuj funkcj `interval_linear_regression_coefficients_estimation()`, kt贸ra bdzie dokonywaa estymacji przedziaowej wartoci wsp贸czynnik贸w r贸wnania regresji liniowej dla danych wejciowych w postaci szeregu szczeg贸owego zgodnie ze schematem z zaczonego do notatnika zestawu wzor贸w.

Opr贸cz zwracanych wartoci granic przedziau funkcja powinna wywietla nastpujcy komunikat:

`a nale偶y do przedziau [X1, Y1] przy zao偶eniu poziomu ufnoci 1- = Z1` </br>
`b nale偶y do przedziau [X2, Y2] przy zao偶eniu poziomu ufnoci 1- = Z2`

gdzie X1, X2, Y1, Y2, Z1 i Z2 s automatycznie uzupeniane przez funkcj, a X1, X2, Y1 i Y2 dodatkowo sformatowane w taki spos贸b, 偶eby wywietlay si z dokadnoci do 4 miejsc po przecinku.

<!-- grade_id: cell-07dc1b0100c687d3, points: 0, grade: false, solution: true -->
```python
def linear_regression_coefficients_estimation(data, alpha):
    """
    Parameters
    -------
    data: DataFrame
        Tablica zawierajca domylny indeks i dwie kolumny "X" i "Y" z wynikami pr贸by losowej.
    alpha: float
        Warto okrelajca poziom istotnoci.
    Returns
    -------
    a_ci: tuple
        Zmienna typu tuple zawierajaca granice przedziau ufnoci parametru a (a_low, a_high).
    b_ci: tuple
        Zmienna typu tuple zawierajaca granice przedziau ufnoci parametru b (b_low, b_high).
    """
    # YOUR CODE HERE
    X = data['X']
    Y = data['Y']
    N = len(X)
    X_mean = X.mean()
    Y_mean = Y.mean()

    sum_a1 = 0
    sum_a2 = 0
    su_sum = 0
    Da_sum = 0
    Db_sum1 = 0
    Db_sum2 = 0

    for x, y in zip(X, Y):
        sum_a1 += (x - X_mean) * (y - Y_mean)
        sum_a2 += (x - X_mean) ** 2

    a = sum_a1 / sum_a2

    b = Y_mean - a * X_mean

    for x, y in zip(X, Y):
        su_sum += (y - (a * x + b)) ** 2

    su = math.sqrt(1 / (N - 2) * su_sum)

    for x, y in zip(X, Y):
        Da_sum += (x - X_mean) ** 2
        Db_sum1 += x ** 2
        Db_sum2 += (x - X_mean) ** 2

    D_a = su / np.sqrt(Da_sum)
    D_b = su * math.sqrt(Db_sum1 / (N * Db_sum2))

    t_val = t.ppf(1 - alpha / 2, N - 2)

    a_low = a - t_val * D_a
    a_high = a + t_val * D_a
    b_low = b - t_val * D_b
    b_high = b + t_val * D_b

    Z1 = 1 - alpha
    Z2 = 1 - alpha

    print(f"a nale偶y do przedziau [{a_low:.4f}, {a_high:.4f}] przy zao偶eniu poziomu ufnoci 1-alpha = {Z1:.4f}")
    print(f"b nale偶y do przedziau [{b_low:.4f}, {b_high:.4f}] przy zao偶eniu poziomu ufnoci 1-alpha = {Z2:.4f}")
    return (a_low, a_high), (b_low, b_high)

print(linear_regression_coefficients_estimation(test_data_2, 0.1))

```

Output:
```
a nale偶y do przedziau [1.9033, 2.0245] przy zao偶eniu poziomu ufnoci 1-alpha = 0.9000
b nale偶y do przedziau [-5.3310, -4.7958] przy zao偶eniu poziomu ufnoci 1-alpha = 0.9000
((1.90330932143819, 2.024457061652292), (-5.331039505544366, -4.79582912649477))
```

<!-- grade_id: cell-54fa1dbb28b96119, points: 8, grade: true, solution: false -->
```python
### Kom贸rka testowa
assert np.all(np.isclose(linear_regression_coefficients_estimation(test_data_2, 0.1), ((1.90330932143819, 2.024457061652292), (-5.331039505544365, -4.7958291264947706))))
print()
assert np.all(np.isclose(linear_regression_coefficients_estimation(test_data_2, 0.05), ((1.8916219551002786, 2.0361444279902035), (-5.382672328053067, -4.744196303986069))))
```

Output:
```
a nale偶y do przedziau [1.9033, 2.0245] przy zao偶eniu poziomu ufnoci 1-alpha = 0.9000
b nale偶y do przedziau [-5.3310, -4.7958] przy zao偶eniu poziomu ufnoci 1-alpha = 0.9000

a nale偶y do przedziau [1.8916, 2.0361] przy zao偶eniu poziomu ufnoci 1-alpha = 0.9500
b nale偶y do przedziau [-5.3827, -4.7442] przy zao偶eniu poziomu ufnoci 1-alpha = 0.9500
```